---
title       : Advanced fitting and inspection
description : Insert the chapter description here
--- type:NormalExercise lang:r xp:100 skills:1 key:520ffecd81
## TODO

This chapter is a placeholder for future content.

*** =instructions

*** =hint

*** =pre_exercise_code
```{r}

```

*** =sample_code
```{r}

```

*** =solution
```{r}

```

*** =sct
```{r}

```
