---
title       : Basics of model syntax, fitting, and inspection
description : Insert the chapter description here
attachments :
  slides_link : https://s3.amazonaws.com/assets.datacamp.com/course/teach/slides_example.pdf

--- type:NormalExercise lang:r xp:100 skills:1 key:2f3ceb3b77
## Model syntax



*** =instructions
- instruction 1

*** =hint
- hint 1

*** =pre_exercise_code
```{r}
# You can also prepare your dataset in a specific way in the pre exercise code
```

*** =sample_code
```{r}
```

*** =solution
```{r}
```

*** =sct
```{r}
# SCT written with testwhat: https://github.com/datacamp/testwhat/wiki
```
